```python
# Initialize Otter
import otter
grader = otter.Notebook("lab_a.ipynb")
```

# Lab A: Data Management

## Setup

```python
# Import some useful functions
from numpy import *
from numpy.random import *
from datascience import *
from statsmodels.formula.api import *

# Define some useful functions
def correlation(array_1, array_2):
    return corrcoef(array_1, array_2).item(1)

# Customize look of graphics
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
plt.rcParams['figure.dpi'] = 60
%matplotlib inline

# Force display of all values 
from IPython.core.interactiveshell import InteractiveShell
InteractiveShell.ast_node_interactivity = "all"

# Hide some unnecessary warning messages
import warnings
warnings.filterwarnings("ignore")
import numpy
numpy.int = numpy.int_
```

## Competing with a Big Box Retailer

### Business Decision

A general retailer that competes with Best Buy wants to determine how well Best Buy is run.  It hopes to craft its own corporate strategy to exploit any Best Buy weakness it might discover. (A company's corporate strategy is its decision about where to apply its limited resources.)

### Data

Retrieve Best Buy data, show the number of observations and the first few observations.

```python
data = ...
...
data
```

```python
grader.check("q1")
```

### Analysis

Show the first year and last year that revenue was recorded.  
Show the time span (in years) over which revenue was recorded. Do not round to nearest integer.

```python
first_year = ...
last_year = ...
span = ...

first_year
last_year
span
```

```python
grader.check("q2_1")
```

Show the quarterly revenue amounts.

```python
revenue = ...
revenue
```

```python
grader.check("q2_2")
```

Calculate the sum of Best Buy's revenue over the time span.

```python
...
```

```python
grader.check("q2_3")
```

<p style="text-align:left; font-size:10px;">
Copyright (c) Huntsinger Associates, LLC
<span style="float:right;">
Document revised August 22, 2023
</span>
</p>

Tests for q1 (points: None)
```
>>> data.num_rows == 83
True
```
```
>>> data.labels == ('Year', 'Quarter', 'Revenue', 'Cost of Goods Sold', 'Gross Profit (million)')
True
```


Tests for q2_1 (points: None)
```
>>> first_year == 1995
True
```
```
>>> last_year == 2015
True
```
```
>>> span == 20.75
True
```


Tests for q2_2 (points: None)
```
>>> len(revenue) == 83
True
```


Tests for q2_3 (points: None)
```
>>> round(sum(revenue)) == 607015
True
```
