```python
# Initialize Otter
import otter
grader = otter.Notebook("lab_n.ipynb")
```

# Lab N: Interaction

## Setup

```python
import builtins
import math
import warnings
from datascience import Table, make_array
from statsmodels.formula.api import ols

# Force display of all values 
from IPython.core.interactiveshell import InteractiveShell
InteractiveShell.ast_node_interactivity = "all"

# Handle some obnoxious warning messages
warnings.filterwarnings("ignore")
```

## Advertising & Pricing Beauty Products

### Business Decision

Tyler Beauty Products advertises and sells beauty products.  It recently conducted a set of marketing programs at various retail stores, where products were priced at 3 different levels: \\$2.00, \\$2.50, or \\$3.00.  Each of these stores advertised at 2 different budget levels: \\$50,000 or \\$100,000.  The product manager must now decide on advertising budgets of \\$50,000 or \\$100,000 at each of several new stores.  For various strategic reasons, each of the new stores has already decided whether to price the beauty products at \\$2.00, \\$2.50, or \\$3.00.

### Retrieve Data

Retrieve data from file 'Tyler.csv'. Show ALL marketing program results.

```python
data = ...
data
```

```python
grader.check("q1")
```

### Analysis

#### Model 1

Build a multiple linear regression model to predict 'Sales' based on 'Price' and 'Budget'.  Show the model goodness of fit (R^2).  Show the model parameter values (intercept and coefficients).

```python
model_formula = ...
model = ...
model.rsquared
model.params
```

```python
grader.check("q2")
```

Build a report that shows what sales and profits the model would predict for each combination of price and budget.
Calculate profit as sales minus budget.

```python
report = Table().with_columns('Price', make_array(2, 2.5, 3))
report = report.with_column('Budget', make_array(50, 50, 50))

sales_predicted_50 = ...
report = report.with_column('sales_predicted_50', sales_predicted_50)

profit_50 = ...
report = report.with_column('profit_50', profit_50)

report = report.with_column('Budget', make_array(100, 100, 100)) # replace Budget amount

sales_predicted_100 = ...
report = report.with_column('sales_predicted_100', sales_predicted_100)

profit_100 = ...
report = report.with_column('profit_100', profit_100)

report = report.drop('Budget')

report
```

```python
grader.check("q3")
```

Visualize predicted profits at budget \$50,000 and budget \$100,000 vs. price, as a line plot.

```python
x_axis_column_name = ...
y_axis_column_name = ...
y_axis_column_name_2 = ...
report.select(x_axis_column_name, y_axis_column_name, y_axis_column_name_2).plot(x_axis_column_name)
```

```python
grader.check("q4")
```

For each price, determine which budget results in the best predicted profit.

```python
for price in make_array(2, 2.5, 3):
    budget = make_array('Budget 50', 'Budget 100')
    profit_50  = ...
    profit_100 = ...
    larger_profit = ...
    (price, budget.item((profit_50<profit_100)*1), larger_profit)
```

```python
grader.check("q5")
```

#### Model 2

Add a synthetic numeric variable called 'priceXbudget' to the data, with values that are the product of 'Price' and
'Budget'. This is an INTERACTION variable. Show ALL the marketing program results.

```python
syth_col_data = ...
data = data.with_column('priceXbudget', syth_col_data)
data.show()
```

```python
grader.check("q6")
```

Build a multiple linear regression model to predict 'Sales' based on 'Price', 'Budget', and 'priceXbudget'.  Show the model goodness of fit (R^2).  Show the model parameter values (intercept and coefficients).

```python
model_formula = ...
model = ols(model_formula, data).fit()
r_squared = ...
model.params
```

```python
grader.check("q7")
```

Build a report that shows what sales and profits the model would predict for each combination of price and budget.
Calculate profit as sales minus budget.

```python
report = Table().with_columns('Price', make_array(2, 2.5, 3))
report = report.with_column('Budget', make_array(50, 50, 50))

array_price_50_budget = ...
report = report.with_column('priceXbudget', array_price_50_budget)

sales_predicted_50 = ...
report = report.with_column('sales_predicted_50', sales_predicted_50)

profit_50 = ...
report = report.with_column('profit_50', profit_50)


report = report.with_column('Budget', make_array(100, 100, 100)) # replace Budget amount

array_price_100_budget = ...
report = report.with_column('priceXbudget', array_price_100_budget)

sales_predicted_100 = ...
report = report.with_column('sales_predicted_100', sales_predicted_100)

profit_100 = ...
report = report.with_column('profit_100', profit_100)

report = report.drop('Budget')
report = report.drop('priceXbudget')

report
```

```python
grader.check("q8")
```

Visualize predicted profits at budget \$50,000 and budget \$100,000 vs. price, as a line plot.

```python
x_axis_column_name_2 = ...
y_axis_column_name_2 = ...
y_axis_column_name_3 = ...
report.select(x_axis_column_name_2, y_axis_column_name_2, y_axis_column_name_3).plot(x_axis_column_name_2)
```

```python
grader.check("q9")
```

For each price, determine which budget results in the best predicted profit.

```python
for price in make_array(2, 2.5, 3):
    budget = make_array('Budget 50', 'Budget 100')
    profit_50  = ...
    profit_100 = ...
    larger_profit = ...
    (price, budget.item((profit_50<profit_100)*1), larger_profit)
```

```python
grader.check("q10")
```

Calculate the price where it becomes more profitable to budget \$50,000 than to budget \$100,000.

$$
\textit{predicted profit at budget 50,000} = \big( i + (c_1)(\textit{price}) + (c_2)(50) + (c_3)(\textit{price})(50) \big) - 50
$$

$$
\textit{predicted profit at budget 100,000} = \big( i + (c_1)(\textit{price}) + (c_2)(100) + (c_3)(\textit{price})(100) \big) - 100
$$

$$
i + (c_1)(\textit{price}) + (c_2)(50) + (c_3)(\textit{price})(50) - 50 = i + (c_1)(\textit{price}) + (c_2)(100) + (c_3)(\textit{price})(100) - 100
$$

$$
(c_2)(50) + (c_3)(\textit{price})(50) - 50 = (c_2)(100) + (c_3)(\textit{price})(100) - 100
$$

$$
(c_2)(50) - (c_2)(100) - 50 + 100 =  (c_3)(\textit{price})(100) - (c_3)(\textit{price})(50)
$$

$$
(c_2)(50) - (c_2)(100) - 50 + 100 = (\textit{price}) \big((c_3)(100) - (c_3)(50)\big)
$$

$$
\cfrac{(c_2)(50) - (c_2)(100) - 50 + 100}{(c_3)(100) - (c_3)(50)}  = \textit{price}
$$


```python
c2 = ...
c3 = ...

crossover_numerator = ...
crossover_denominator = ...
crossover = crossover_numerator / crossover_denominator
crossover
```

```python
grader.check("q11")
```

<p style="text-align:left; font-size:10px;">
Copyright (c) Huntsinger Associates, LLC
<span style="float:right;">
Document revised November 24, 2023
</span>
</p>

Tests for q1 (points: None)
```
>>> data.labels == ('Store', 'Price', 'Budget', 'Sales') and data.num_rows == 24
True
```


Tests for q10 (points: None)
```
>>> any(profit_100 == report.column('profit_100'))
True
```
```
>>> any(profit_100 == report.column('profit_100'))
True
```
```
>>> any(profit_50 == report.column('profit_50'))
True
```
```
>>> any(larger_profit == report.column('profit_100'))
False
```


Tests for q11 (points: None)
```
>>> c2 == model.params[2]
True
```
```
>>> c3 == model.params[3]
True
```
```
>>> crossover_numerator == c2 * 50 - c2 * 100 - 50 + 100
True
```
```
>>> crossover_denominator == c3 * 100 - c3 * 50
True
```


Tests for q2 (points: None)
```
>>> 'Budget' in model_formula
True
```
```
>>> math.isclose(model.rsquared, 0.791, abs_tol=0.001)
True
```
```
>>> hasattr(model, 'rsquared') and hasattr(model, 'params') and (len(model.params) == 3)
True
```


Tests for q3 (points: None)
```
>>> report = report.with_column('Budget', make_array(50, 50, 50))
>>> all(sales_predicted_50 == model.predict(report))
True
```
```
>>> report = report.with_column('Budget', make_array(100, 100, 100))
>>> all(sales_predicted_100 == model.predict(report))
True
```
```
>>> report = report.with_column('Budget', make_array(100, 100, 100))
>>> all(profit_100 == sales_predicted_100 - report.column('Budget'))
True
```
```
>>> report = report.with_column('Budget', make_array(50, 50, 50))
>>> all(profit_50 == sales_predicted_50 - report.column('Budget'))
True
```


Tests for q4 (points: None)
```
>>> x_axis_column_name == 'Price'
True
```
```
>>> y_axis_column_name == 'profit_50'
True
```
```
>>> x_axis_column_name == 'Price'
True
```
```
>>> y_axis_column_name_2 == 'profit_100'
True
```


Tests for q5 (points: None)
```
>>> any(profit_100 == report.column('profit_100'))
True
```
```
>>> any(profit_100 == report.column('profit_100'))
True
```
```
>>> any(profit_50 == report.column('profit_50'))
True
```
```
>>> any(larger_profit == report.column('profit_100'))
True
```


Tests for q6 (points: None)
```
>>> all(syth_col_data == data.column('Price') * data.column('Budget'))
True
```


Tests for q7 (points: None)
```
>>> 'priceXbudget' in model_formula
True
```
```
>>> import math
>>> math.isclose(r_squared, 0.97, abs_tol=0.05)
True
```


Tests for q8 (points: None)
```
>>> report = report.with_column('Budget', make_array(50, 50, 50))
>>> report = report.with_column('priceXbudget', make_array(50 * 2, 50 * 2.5, 50 * 3))
>>> all(sales_predicted_50 == model.predict(report))
True
```
```
>>> report = report.with_column('Budget', make_array(100, 100, 100))
>>> report = report.with_column('priceXbudget', make_array(100 * 2, 100 * 2.5, 100 * 3))
>>> all(sales_predicted_100 == model.predict(report))
True
```


Tests for q9 (points: None)
```
>>> x_axis_column_name_2 == 'Price'
True
```
```
>>> y_axis_column_name_2 == 'profit_50'
True
```
```
>>> y_axis_column_name_3 == 'profit_100'
True
```
