```python
# Initialize Otter
import otter
grader = otter.Notebook("lab_c1.ipynb")
```

# Lab C: Weighted Average

## Setup

```python
# Import some useful functions
import warnings
from datascience import Table
from numpy import mean, median, array

# Force display of all values 
from IPython.core.interactiveshell import InteractiveShell
InteractiveShell.ast_node_interactivity = "all"

# Hide some unnecessary warning messages
warnings.filterwarnings("ignore")
```

## Airline Reservations

### Business Decision

The Manager of Travel Operations is considering setting a policy to encourage his company's executives to do business travel on either of two specific airlines.  He needs to account for how the airlines are suited to the kind of business travel typical of the executives and the airlines' prices.

### Data

The dataset comprises a random sample of recent roundtrip tickets purchased from 2 airlines for travel between New York’s John F. Kennedy Airport (JFK) and Los Angeles Airport (LAX). Each ticket is described by the following variables:

* `carrier`: the airline
* `rtrip_fare`: round trip fare
* `advance`: number of days in advance of the flight that ticket was purchased
* `busclass`: indicator for business class ('1' if business class, '0' if not)
* `nonrefundable`: indicator for whether a ticket was nonrefundable ('1' if nonrefundable, '0' if refundable)
* `satstayover`: indicator for whether a roundtrip ticket includes a Saturday stayover in the destination location

Retrieve the airline ticket dataset, show the number of observations and the first few observations.

```python
ticket_prices = ...
...
ticket_prices
```

```python
grader.check("q1")
```

### Analysis | Weighted Averages of Groups

Calculate the mean and median prices.

```python
mean_ticket_price = ...
median_ticket_price = ...
```

```python
grader.check("q2_1")
```

For each airline, calculate the number of flights.

```python
airline_flights = ...
airline_flights
```

```python
grader.check("q2_2")
```

Calculate statistics to show whether one airline is much cheaper than the other, on average.

```python
airline_mean_prices = ...
airline_mean_prices
```

```python
grader.check("q2_3")
```

Calculate statistics to show whether flights with Saturday stayovers are much cheaper than other flights, on average.

```python
stayover_groups = ticket_prices.group('satstayover', mean)
stayover_groups

diff_satstayover = stayover_groups.column('rtrip_fare mean').item(0) - stayover_groups.column('rtrip_fare mean').item(1)
diff_satstayover
```

```python
grader.check("q2_4")
```

<p style="text-align:left; font-size:10px;">
Copyright (c) Huntsinger Associates, LLC
<span style="float:right;">
Document revised September 1, 2023
</span>
</p>

Tests for q1 (points: None)
```
>>> ticket_prices.num_rows == 10266
True
```
```
>>> ticket_prices.labels == ('carrier', 'rtrip_fare', 'advance', 'busclass', 'nonrefundable', 'satstayover')
True
```


Tests for q2_1 (points: None)
```
>>> round(mean_ticket_price) == 369
True
```
```
>>> round(median_ticket_price) == 328
True
```


Tests for q2_2 (points: None)
```
>>> all(airline_flights.labels == array(['carrier', 'count']))
True
```
```
>>> all(airline_flights.column('carrier') == ['AA', 'UA'])
True
```


Tests for q2_3 (points: None)
```
>>> all(airline_mean_prices.labels == array(['carrier', 'rtrip_fare mean', 'advance mean', 'busclass mean', 'nonrefundable mean', 'satstayover mean']))
True
```
```
>>> all(airline_mean_prices.column('carrier') == ['AA', 'UA'])
True
```


Tests for q2_4 (points: None)
```
>>> all(stayover_groups.column('satstayover') == [0, 1])
True
```
```
>>> round(diff_satstayover) == 150
True
```
