```python
# Initialize Otter
import otter
grader = otter.Notebook("lab_c3.ipynb")
```

# Lab C: Weighted Average

## Setup

```python
import warnings
from datascience import Table, make_array
from numpy import mean, unique

# Force display of all values 
from IPython.core.interactiveshell import InteractiveShell
InteractiveShell.ast_node_interactivity = "all"

# Hide some unnecessary warning messages
warnings.filterwarnings("ignore")
```

## Restaurant Reviews

### Business Decision

An independent marketing firm recently surveyed some of the city's restaurant patrons for their overall ratings of several restaurants.  It is not clear how well the surveyed patrons represent the city's entire population of patrons.  The owner of restaurant \#135060 is considering launching a marketing campaign that would tout its favorable customer ratings.

### Data

The dataset comprises several ratings.  Each observation corresponds to a customer's review of a restaurant.

Retrieve the restaurant survey dataset, show the number of observations and the first few observations

```python
restaurant_reviews = ...
...
restaurant_reviews
```

```python
grader.check("q1_1")
```

Show all the possible rating levels.

```python
rating_levels = ...
rating_levels
```

```python
grader.check("q1_2")
```

### Analysis | Weighted Average by Group, missing data, proportions unknown

Build and show a table comprising reviews for restaurant #135060.

```python
restaurant1 = ...
restaurant1.num_rows
restaurant1.show()
```

```python
grader.check("q2_1")
```

Calculate the mean rating for restaurant #135060.

```python
avg = ...
avg
```

```python
grader.check("q2_2")
```

Calculate the number of student customers for restaurant #135060.  
Calculate the number of professional customers for restaurant #135060.  
Calculate the number of unemployed customers for restaurant #135060.  
Calculate the number of customers with known activity type for restaurant #135060.

```python
# Calculate the number of student customers for restaurant #135060
n_student = ...
n_student

# Calculate the number of professional customers for restaurant #135060
n_professional = ...
n_professional

# Calculate the number of unemployed customers for restaurant #135060
n_unemployed = ...
n_unemployed

# Calculate the number of customers with known activity type for restaurant #135060
n = ...
n
```

```python
grader.check("q2_3")
```

Calculate the mean rating by student customers for restaurant #135060.  
Calculate the mean rating by professional customers for restaurant #135060.  
Calculate the mean rating by unemployed customers for restaurant #135060.

```python
# Calculate the mean rating by student customers for restaurant #135060
mean_student = ...
mean_student

# Calculate the mean rating by professional customers for restaurant #135060
mean_professional = ...
mean_professional

# Calculate the mean rating by unemployed customers for restaurant #135060
mean_unemployed = ...
mean_unemployed
```

```python
grader.check("q2_4")
```

Calculate the weighted average rating by activity type for restaurant #135060,
assuming activity types for the population are equally distributed.

```python
avg_equal = ...
avg_equal
```

```python
grader.check("q2_5")
```

Calculate the weighted average rating by activity type for restaurant #135060,
assuming activity types for the population are distributed as they are for restaurant #135060.

```python
avg_sample = ...
avg_sample
```

```python
grader.check("q2_6")
```

Calculate the weighted average rating by activity type for restaurant #135060,
assuming activity types for the population are distributed as in all the data.

```python
n_all_student = ...
n_all_professional = ...
n_all_unemployed = ...
n_all = ...

n_all_student
n_all_professional
n_all_unemployed
n_all

avg_all = ...
avg_all
```

```python
grader.check("q2_7")
```

Build a table to compare rating averages calculated various ways like this:
* One column for calculation method name + one column for distribution assumption + one column for rating averages
* Four rows for calculation methods
  * First row is mean
  * Second row is weighted average by activity, assuming population distribution is equal
  * Third row is weighted average by activity, assuming population distribution is as in restaurant sample
  * Fourth row is weighted average by activity, assuming population distribution is as in all data

```python
row_values = ...

rating_averages = Table().with_columns('Method', make_array('mean', 'weighted average by activity', 'weighted average by activity', 'weighted average by activity'),
                     'Distribution Assumption', make_array('', 'equal', 'as in restaurant sample', 'as in all data'),
                     'Restaurant #135060', row_values)
```

```python
grader.check("q2_8")
```

<p style="text-align:left; font-size:10px;">
Copyright (c) Huntsinger Associates, LLC
<span style="float:right;">
Document revised September 1, 2023
</span>
</p>

Tests for q1_1 (points: None)
```
>>> restaurant_reviews.num_rows == 1161
True
```
```
>>> len(restaurant_reviews.columns) == 23
True
```


Tests for q1_2 (points: None)
```
>>> all(rating_levels == [0, 1, 2])
True
```


Tests for q2_1 (points: None)
```
>>> restaurant1.num_rows == 22
True
```
```
>>> all(restaurant1['placeID'] == 135060)
True
```


Tests for q2_2 (points: None)
```
>>> round(avg, 2) == 1.14
True
```


Tests for q2_3 (points: None)
```
>>> n_student == 17
True
```
```
>>> n_professional == 3
True
```
```
>>> n_unemployed == 1
True
```
```
>>> n == 21
True
```


Tests for q2_4 (points: None)
```
>>> round(mean_student, 2) == 1.18
True
```
```
>>> round(mean_professional, 2) == 1.33
True
```
```
>>> mean_unemployed == 0.0
True
```


Tests for q2_5 (points: None)
```
>>> round(avg_equal, 2) == 0.84
True
```


Tests for q2_6 (points: None)
```
>>> round(avg_sample, 2) == 1.14
True
```


Tests for q2_7 (points: None)
```
>>> n_all_student == 971
True
```
```
>>> n_all_professional == 119
True
```
```
>>> n_all_unemployed == 14
True
```
```
>>> n_all == 1104
True
```
```
>>> round(avg_all, 2) == 1.18
True
```


Tests for q2_8 (points: None)
```
>>> all(row_values == make_array(avg, avg_equal, avg_sample, avg_all))
True
```
